```python
from grading_tools import check, auto_marking_message
import numpy as np
from math import *
```

Output:
```
<IPython.core.display.HTML object>
```

# Exercises

## Exercise

Fill in the blanks.

If $\Sigma$ is a positive definite symmetric matrix then a matrix satisfying $A A^\top=\Sigma$, is called a
p---- s----- r--- of $\Sigma$. The C------- d------------ of $\Sigma$ is the unique
p---- s----- r--- of $\Sigma$ with positive diagonal.

<!-- grade_id: c94260, points: 0, grade: false, solution: true -->
```python
ans1 = "p----- s----- r---"
ans2 = "C------- d------------"
ans3 = "p----- s----- r---"
#BEGIN SOLUTION
ans1 = "pseudo square root"
ans2 = "Cholesky decomposition"
ans3 = ans1
#END SOLUTION
```

<!-- grade_id: 9510c7, points: 1, grade: true, solution: false -->
```python
check( '9510c7.1', ans1)
check( '9510c7.2', ans2)
check( '9510c7.3', ans3)
auto_marking_message()
```

Output:
```
Auto marking message: 😍 Correct
```

## Exercise

Prove the following:

**Lemma:** $W_t$ is a discrete time Wiener process if and only if $Z_0 + \tilde{\mu} t + \sigma W_t$ is a Browian motion
with drift $\tilde{\mu}$ and volatility $\sigma$. If $Z_t$ is a discrete time Brownian motion, then $\frac{1}{\sigma}(Z_t-Z_0-\tilde{\mu} t)$ is a discrete
time Wiener process.

## Exercise

The following claim is made in the notes. Prove it:
    
"So by the properties of Cholesky decomposition, if we define    
$$
\left(
    \begin{array}{c}
    \epsilon_t^\prime \\
    \epsilon_\frac{t}{2}^\prime
    \end{array}
\right)
:= A^{-1}
\left(
    \begin{array}{c}
    W_t \\
    W_\frac{t}{2}
    \end{array}
\right)
$$
then $\epsilon_t^\prime$ and $\epsilon_\frac{t}{2}^\prime$ are independent standard normal random
variables."

### Solution

Write $\Sigma$ for the covariance matrix of $W_t$ and $W_{\frac{t}{2}}$.
Since $A$
is the Cholesky decomposition of $\Sigma$ it is a pseudo-square root. Hence
$$A A^\top = \Sigma$$
Since $A$ is lower triangular with positive diagonal, it is invertible (as it is already in echelon form). Hence
$$1 = A^{-1} \Sigma (A^{-1})^\top$$
(I'm using the identity $(A^{-1})^\top=(A^\top)^{-1}$ here too)
 
Using vector notation in the obvious way, we have:
$\mathbf{\epsilon}^\prime:=A^{-1} \mathbf{W}$
We can write down (because we've done this type of exercise before) that the covarince matrix of $\epsilon$ is
$$A^{-1} \Sigma (A^{-1})^\top$$

## Exercise

Suppose that $W_t$ is a discrete time Wiener process defined on ${\cal T}=\{0,\frac{1}{4},\frac{1}{2},\frac{1}{4},1\}$ compute
the covariance matrix of the vector

$$(W_1,W_{\frac{1}{2}},W_{\frac{1}{4}},W_{\frac{3}{4}}).$$

Store the answer in a matrix called `cov_matrix`. What is the Cholesky decomposition of this matrix? Store the answer in a matrix called `chol_matrix`

We may simulate $(W_1,W_\frac{1}{2},W_\frac{1}{4}, W_\frac{3}{4})$ using

$$
\left(
\begin{array}{c}
W_1 \\
W_\frac{1}{2} \\
W_\frac{1}{4} \\
W_\frac{3}{4}
\end{array}
\right)
=
\sqrt{\frac{1}{4}}
\left(
\begin{array}{cccc}
1 & 1 & 1 & 1 \\
1 & 1 & 0 & 0 \\
1 & 0 & 0 & 0 \\
1 & 1 & 1 & 0
\end{array} \right)
\left(
\begin{array}{c}
\epsilon_\frac{1}{4}^\prime \\
\epsilon_\frac{1}{2}^\prime \\
\epsilon_\frac{3}{4}^\prime \\
\epsilon_1^\prime
\end{array}
\right)
$$
for some independent standard normals $\epsilon^\prime_t$ determining the increments of $W_t$.

Hence the covariance matrix of $(W_\frac{1}{4},W_\frac{1}{2},W_\frac{3}{4}, W_1)$ is
$$A 1_n A^\top$$
where
$$A=
\sqrt{\frac{1}{4}}
\left(
\begin{array}{cccc}
1 & 1 & 1 & 1 \\
1 & 1 & 0 & 0 \\
1 & 0 & 0 & 0 \\
1 & 1 & 1 & 0
\end{array} \right)
$$
and $1_n$ denotes the identity matrix.

<!-- grade_id: c57c79, points: 0, grade: false, solution: true -->
```python
# BEGIN SOLUTION
A = 0.5*np.array([[1,1,1,1],[1,1,0,0],[1,0,0,0],[1,1,1,0]])
cov_matrix = A @ A.transpose()
chol_matrix = np.linalg.cholesky(cov_matrix)
# END SOLUTION
```

<!-- grade_id: 855c7a, points: 1, grade: true, solution: false -->
```python
for i in range(0,4):
    for j in range(0,4):
        check('855c7a:'+str(i)+str(j), cov_matrix[i,j])
for i in range(0,4):
    for j in range(0,4):
        check('855c7b:'+str(i)+str(j), chol_matrix[i,j])
auto_marking_message()
```

Output:
```
Auto marking message: ✔ Correct
```

## Exercise

Let $n=2$. Write a Python function `simulate_wiener_4_points` that takes as input an $N \times 4$ matrix `epsilon` of independent,
standard random normal variables and returns a matrix with each row representing a different path of the Wiener process and columns representing the time points $(0,\frac{1}{4},\frac{1}{2},\frac{3}{4},1)$ in sequence.
You must follow Wiener's construction so that you simulate times in the order $\{W_0, W_1, W_\frac{1}{2}, W_\frac{1}{4}, W_\frac{3}{4}, \ldots\}$.

Check numerically that the variances of the differences $\delta W_i:= W_{\frac{i}{2^n}}-W_{\frac{i-1}{2^n}}$ take
the value you expect.

Check numerically that the differences $\delta W_{i+1}$ and $\delta W_{j+1}$ are uncorrelated if $i\neq j$.



<!-- grade_id: cdbc7a, points: 0, grade: false, solution: true -->
```python
def simulate_wiener_4_points( epsilon ):
    ### BEGIN SOLUTION
    n_samples = epsilon.shape[0]
    chol = chol_matrix
    r = chol @ epsilon.transpose()
    w_1 = r[0,:]
    w_half = r[1,:]
    w_quarter = r[2,:]
    w_three_quarters = r[3,:]
    z = np.zeros( n_samples )
    return np.array([ z, w_quarter, w_half, w_three_quarters, w_1 ]).transpose()
    ### END SOLUTION
```

<!-- grade_id: 8f7b25, points: 1, grade: true, solution: false -->
```python
# Check the shape
n_samples=10
W = simulate_wiener_4_points(np.random.randn(n_samples, 4))
assert W.shape==(10,5)

# Check that you are using Wiener's construction - this means that
# W_1 will be determined by the first column of n_samples.
v1 = 3
v2 = 2
test_array = np.array([[v1,0,0,0],[v2,0,0,0]])
W = simulate_wiener_4_points(test_array)
assert abs(W[0,0]-0)==0
assert abs(W[0,4]-v1)==0
assert abs(W[1,0]-0)==0
assert abs(W[1,4]-v2)==0

# The tests you have been asked to perform yourself will also be run, but are hidden.

### BEGIN HIDDEN TESTS
n = 2
n_samples=100000
W = simulate_wiener_4_points(np.random.randn(n_samples, 4))
delta_W = np.zeros((n_samples, 4))
for i in range(0,2**n):
    delta_W[:,i] = W[:,i+1]-W[:,i]
    
delta_cov = np.cov( delta_W, rowvar=False )
expected = 0.25*np.identity(4)
for i in range(0,4):
    for j in range(0,4):
        assert abs( delta_cov[i,j]-expected[i,j])<0.03
### END HIDDEN TESTS

auto_marking_message()
```

Output:
```
Auto marking message: 😊 Correct
```

```python

```