```python
from grading_tools import check, auto_marking_message
```

Output:
```
<IPython.core.display.HTML object>
```

# Exercises

## Exercise

What is the covariance matrix of  $(W_\frac{1}{4},W_\frac{1}{2},W_\frac{3}{4},W_{1})$? Use this
to write a unit test for the function `wiener_interpolate`. You must call your unit test `test_wiener_interpolate`.

<!-- grade_id: 066525, points: 0, grade: false, solution: true -->
```python
from fmlib import wiener_interpolate, one_step_wiener
import numpy as np
np.random.seed(0)

def test_wiener_interpolate():
    ### BEGIN SOLUTION
    n_samples = 100000
    sample = np.zeros((n_samples,4))
    for i in range(0,n_samples):
        W, t = one_step_wiener(1)
        W, t = wiener_interpolate(W,t)
        W, t = wiener_interpolate(W,t)
        sample[i,:]=W[1:]
    c = np.cov( sample, rowvar = False)
    A = 0.5*np.array([[1,0,0,0],[1,1,0,0],[1,1,1,0],[1,1,1,1]])
    expected_c = A @ A.transpose()
    for i in range(0,3):
        for j in range(0,3):
            assert abs(c[i,j]-expected_c[i,j])<0.03
    ### END SOLUTION
```

<!-- grade_id: 93cd19, points: 1, grade: true, solution: false -->
```python
test_wiener_interpolate() # Your test should pass!
### BEGIN HIDDEN TESTS
# We now redefine Wiener interpolate and check that the test fails
def wiener_interpolate():
    halfway_times = compute_halfway_times(t)
    halfway_values = simulate_intermediate_values(W, t)
    new_W = 0.8*riffle(W,halfway_values) # DELIBERATE ERROR!
    new_t = riffle(t,halfway_times)
    return new_W, new_t

passed = False
try:
    test_wiener_interpolate()
    passed = True
except:
    pass

if passed:
    raise Exception("Your test didn't notice that the function no longer works")
### END HIDDEN TESTS
auto_marking_message()
```

Output:
```
Auto marking message: ✨ Correct
```

```python

```

## Exercise

Write a function `simulate_multivariate_normal` that can simulate `n_samples` samples from a multivariate normal distribution with a given covariance
matrix $\Sigma$ and with mean $0$. Each row should represent a different sample. Write a unit test for this function.

<!-- grade_id: 83feb1, points: 0, grade: false, solution: true -->
```python
def simulate_multivariate_normal( n_samples, sigma ):
    ### BEGIN SOLUTION
    n = sigma.shape[0]
    L = np.linalg.cholesky( sigma )
    epsilon = np.random.randn(n, n_samples)
    return (L @ epsilon).transpose()
    ### END SOLUTION

def test_simulate_multivariate_normal():
    ### BEGIN SOLUTION
    dr_A_test_simulate_multivariate_normal() # Defined below!
    ### END SOLUTION
```

<!-- grade_id: e672e6, points: 1, grade: true, solution: false -->
```python
sigma = np.array([[1,0.5],[0.5,1]])
ans = simulate_multivariate_normal(10, sigma)
assert ans.shape==(10,2)

### BEGIN HIDDEN TESTS
def dr_A_test_simulate_multivariate_normal():
    sigma = np.array([[1,0.5,0.3],[0.5,2,0],[0.3,0,3]])
    r = simulate_multivariate_normal( 100000, sigma )
    c = np.cov(r,rowvar=False)
    for i in range(0,3):
        for j in range(0,3):
            assert abs( c[i,j]-sigma[i,j])<0.05

# Check your function works
dr_A_test_simulate_multivariate_normal()

# Check your test fails by changing the function
# to one that doesn't work
old_f = simulate_multivariate_normal

def simulate_multivariate_normal( n_samples, sigma ):
     return 0.5*old_f( n_samples, sigma )

passed = False
try:
    test_wiener_interpolate()
    passed = True
except:
    pass

if passed:
    raise Exception("Your test didn't notice that the function no longer works")

# Switch the function back to a working version
simulate_multivariate_normal = old_f
### END HIDDEN TESTS

test_simulate_multivariate_normal()
auto_marking_message()
```

Output:
```
Auto marking message: 😀 Correct
```

```python

```