# Exercises

```python
import fmlib
import numpy as np
```

Output:
```
<IPython.core.display.HTML object>
```

## Exercise

Use Ito's Lemma to show that $X_t=W_t^3$ solves an SDE of the form
$$d X_t = a W_t dt + b W_t^c dW_t$$
for constants $a$, $b$ and $c$ that you should determine. 
Write the code needed to simulate a vector containing the values of $X$
obtained from the Euler Maruyama scheme for this SDE given the values of W
at the time time points t.

<!-- grade_id: c90d03, points: 0, grade: false, solution: true -->
```python
def simulate_X_from_euler_maruyama(W,t):
    # BEGIN SOLUTION
    X = np.zeros(len(t))
    X[0] = 0
    for i in range(0,len(t)-1):
        dt = t[i+1]-t[i]
        dW = W[i+1]-W[i]
        X[i+1]=X[i] + 3*W[i] * dt + 3*W[i]**2*dW
    return X
    # END SOLUTION
```

<!-- grade_id: 25fa8d, points: 1, grade: true, solution: false -->
```python
np.random.seed(0)
W,t = fmlib.one_step_wiener(1)
W,t = fmlib.wiener_interpolate(W,t,20)
actual_X = simulate_X_from_euler_maruyama(W,t)
expected_X = W**3
assert abs( expected_X[-1] - actual_X[-1] )<0.01
### BEGIN HIDDEN TESTS
import grading_tools
grading_tools.check('25fa8d',actual_X)
grading_tools.auto_marking_message()
### END HIDDEN TESTS
```

Output:
```
Auto marking message: ✔ Correct
```

## Exercise

Use Ito's Lemma to prove that if $Z_t$ satisfies
$$
d Z_t = (\mu - \frac{1}{2}\sigma^2)\, d t + \sigma \, d W_t
$$
for constants $\mu$ and $\sigma$ then if we define $S_t=\exp(Z_t)$ we will have
$$
d S_t = S_t(\mu \, dt + \sigma \, d W_t ).
$$

Deduce that the solution of
$$
d S_t = S_t(\mu \, dt + \sigma \, d W_t ).
$$
given $S_0$ is given by
$$
S_t = S_0 e^{(\mu - \frac{1}{2}\sigma^2)t + \sigma W_t}
$$

### Solution

Take $f(z)=\exp(z)$. Plugging this into Ito's Lemma

$$
\begin{split}
d f(Z)_t &= f^\prime(Z_t) (\mu-\tfrac{1}{2}\sigma^2) \, dt + \tfrac{1}{2} f^{\prime\prime}(Z_t) \sigma^2 \, dt + f^{\prime}(Z_t) \sigma \, dW_t \\
&= \exp(z) (\mu \, dt + \sigma \, dW_t) \\
&= f(Z_t) (\mu \, dt + \sigma \, dW_t)
\end{split}
$$

hence

$$
\begin{split}
d S_t &= S_t (\mu \, dt + \sigma \, dW_t)
\end{split}
$$

But $Z_t=Z_0 + (\mu - \frac{1}{2}\sigma^2)t + \sigma W_t$ by solving the SDE for $Z_t$ (recall we showed how to solve constant coefficient SDEs last week).
Hence $Z_t=\log(S_0) + (\mu - \frac{1}{2}\sigma^2)t + \sigma W_t$ and so
$$
S_t=S_0 e^{ (\mu - \frac{1}{2}\sigma^2)t + \sigma W_t}
$$


```python

```