# Exercises

## Exercise

Write a function to price a digital option with strike $K$ in the Black Scholes model by the Monte Carlo method.
Your function should return two values, lower and upper, giving a $99.999%$ confidence interval for the price.
The hidden tests will fail if the true price is not in your confidence interval (in which case you would be unlucky)

<!-- grade_id: ba816f, points: 0, grade: false, solution: true -->
```python
import fmlib
import numpy as np
import scipy
from grading_tools import check, auto_marking_message

def price_digital_call_by_monte_carlo( S0, r, sigma, K, T, n_steps, n_paths ):
    ### BEGIN SOLUTION
    drift = r
    S_twiddle, times = fmlib.simulate_gbm_paths(S0, drift, sigma, T, n_steps, n_paths )
    S_T = S_twiddle[:,-1]
    payoff = S_T>K
    price = np.exp(-r*T)*np.mean( payoff )

    sigma_sample = np.std( np.exp(-r*T) * payoff )
    p = 99.999
    alpha = scipy.stats.norm.ppf((1-p/100)/2)
    print( alpha )
    lower = price + alpha*sigma_sample/np.sqrt(n_paths)
    upper = price - alpha*sigma_sample/np.sqrt(n_paths)
    return lower, upper
    ### END SOLUTION
```

Output:
```
<IPython.core.display.HTML object>
```

<!-- grade_id: 3815fd, points: 1, grade: true, solution: false -->
```python
S0 = 100
r = 0.05
sigma = 0.2
K = 105
T = 2
low, high=price_digital_call_by_monte_carlo(S0,r,sigma,K,T,1,1000000)
assert low<high
auto_marking_message()
```

Output:
```
-4.417173413467605
Auto marking message: 😻 Correct
```

<!-- grade_id: 6a2b27, points: 1, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
def black_scholes_digital_call_price(S,r,sigma,K,T):
    d2 = 1/(sigma*np.sqrt(T))*(np.log(S/K)+(r-0.5*sigma**2)*T)
    return scipy.stats.norm.cdf(d2)*np.exp(-T*r)

bs = black_scholes_digital_call_price(S0,r,sigma,K,T)
assert low<bs
assert bs<high
# check the confidence interval is in the right ballpark, it's actually about 0.004
assert high-low<0.005
assert high-low>0.003
### END HIDDEN TESTS
```

## Exercise

True or False? As you increase the number of steps in the Monte Carlo simulation to price the digital option you increase the accuracy of the computation of the option price.

<!-- grade_id: 83391d, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
False
### END SOLUTION
```

Output:
```
False
```

```python
check('83391d',_)
auto_marking_message()
```

Output:
```
Auto marking message: 👍 Correct
```

### Solution

It is False. We can simulate $S_T$ perfectly using only one time step as have solved the SDE.