```python
import numpy as np
import matplotlib.pyplot as plt
import fmlib
import sympy as sp
from grading_tools import check, auto_marking_message
```

Output:
```
<IPython.core.display.HTML object>
```

# Exercises

## Exercise

Use `sympy` to compute

$$\frac{\partial^2}{\partial x \, \partial y} e^{\sin(x)\cos(y)}$$

Store your answer in a variable called `answer`.

<!-- grade_id: 7ff21f, points: 0, grade: false, solution: true -->
```python
x, y = sp.symbols('x y')
# BEGIN SOLUTION
answer = sp.diff( sp.exp(sp.sin(x)*sp.cos(y)), x, y )
# END SOLUTION
```

<!-- grade_id: f79702, points: 1, grade: true, solution: false -->
```python
check('f79702',answer.subs([(x,0.1),(y,0.2)]).evalf())
auto_marking_message()
```

Output:
```
Auto marking message: 😊 Correct
```

## Exercise

Use `sympy` to compute the value of 

$$
\int_{-\infty}^{\infty} \frac{1}{\sqrt{2 \pi}} x^4 e^{-\frac{x^2}{2}} \, d x.
$$

Store your answer in a variable called answer.
You will probably need to use the documentation of `sympy` to complete this task.

<!-- grade_id: e779fb, points: 0, grade: false, solution: true -->
```python
answer = sp.integrate(1/sp.sqrt(2*sp.pi)*x**4*sp.exp(-x**2/2), (x,-sp.oo,sp.oo))
```

<!-- grade_id: 90ceae, points: 1, grade: true, solution: false -->
```python
check( '90ceae', answer )
auto_marking_message()
```

Output:
```
Auto marking message: 😺 Correct
```

## Exercise

Use the `apply_itos_lemma` function to show that if 
$$
d S_t = S_t( \mu dt + \sigma dW_t)
$$
then this implies
$$
d \log(S)_t = (\mu - \frac{1}{2}\sigma^2) dt + \sigma dW_t.
$$
Repeat this calculation by hand.

This question is not automatically marked.


```python
# Solution

S, mu, sigma = sp.symbols(['S','mu','sigma'])
S_vec = sp.Matrix([S])
a = sp.Matrix([S*mu])
b = sp.Matrix([S*sigma])
f = sp.Matrix([sp.log(S)])

c_dt, c_dW = fmlib.apply_itos_lemma(S_vec,a,b,f)
assert( c_dt==sp.Matrix([mu-sigma**2/2]))
assert( c_dW==sp.Matrix([sigma]) )
```

## Exercise

Write a function `euler_maruyama` that takes as parameters the initial condition $X_0$, the coefficient functions $a$ and $b$ and
vectors `W`, `times` containing the noise process and the time and returns a vector containing all the corresponding values of $X$
if one applies the Euler Maruyama scheme to the equation
$$
d X_t = a(X_t,t) \, d t + b(X_t,t) \, d W_t.
$$
Your code should be able to cope with any Python functions $a(X,t)$ and $b(X,t)$. The test checks that your code works for geometric Brownian motion.

<!-- grade_id: b01a30, points: 0, grade: false, solution: true -->
```python
def euler_maruyama(X0,a,b,W,times):
    ### BEGIN SOLUTION
    X = np.zeros(len(W))
    X[0] = X0
    for i in range(0,len(times)-1):
        t = times[i]
        dt = times[i+1]-times[i]
        dW = W[i+1]-W[i]
        X[i+1]=X[i] + a(X[i],t)*dt + b(X[i],t)*dW
    return X
    ### END SOLUTION

```

<!-- grade_id: 010ada, points: 1, grade: true, solution: false -->
```python
# Test with geometric Brownian motion
np.random.seed(0)
T = 1
W, times = fmlib.one_step_wiener(T)
W,times = fmlib.wiener_interpolate(W,times,20)

mu = 0.5
sigma = 1
def a_gbm(S,t):
    return mu*S
def b_gbm(S,t):
    return sigma*S

WT = W[-1]
S0 = 100
S = euler_maruyama(S0,a_gbm,b_gbm,W,times);
expectedST = S0*np.exp((mu- 0.5*sigma**2)*T + sigma*WT)
ST = S[-1]
assert abs(expectedST - ST)<0.1 ;
auto_marking_message()
```

Output:
```
Auto marking message: 👍 Correct
```

```python

```